```python
# Initialize Otter
import otter
grader = otter.Notebook("hw9.ipynb")
```

# Практическое задание

1. Напишите класс Person, который будет содержать атрибуты `name` и `age`, и метод `introduce`, который будет возвращать информацию о человеке в формате "Меня зовут `[name]`, мне `[age]` лет."

```python
# BEGIN SOLUTION NO PROMPT
class Person:
    def __init__(self, name, age):
        self.name = name
        self.age = age

    def introduce(self):
        return f"Меня зовут {self.name}, мне {self.age} лет."

person = Person("Иван", 25)
intro = person.introduce()


# END SOLUTION
""" # BEGIN PROMPT
class Person:
    def __init__(self, name, age):
        # ТУТ ТВОЙ КОД

    def introduce(self):
       # ТУТ ТВОЙ КОД

person = Person("Иван", 25)
person.introduce()
""" # END PROMPT
```

Output:
```
' # BEGIN PROMPT\nclass Person:\n    def __init__(self, name, age):\n        # ТУТ ТВОЙ КОД\n\n    def introduce(self):\n       # ТУТ ТВОЙ КОД\n\nperson = Person("Иван", 25)\nperson.introduce()\n'
```

```python
grader.check("q1")
```

2.  Создайте класс `Rectangle`, который принимает два аргумента в инициализаторе: `width` (ширина) и `height` (высота). Добавьте метод `area`, который будет возвращать площадь прямоугольника.

```python
# BEGIN SOLUTION NO PROMPT
class Rectangle:
    def __init__(self, width, height):
        self.width = width
        self.height = height

    def area(self):
        return self.width * self.height


# END SOLUTION
""" # BEGIN PROMPT
class Rectangle:
    def __init__(self, width, height):
        # ТУТ ТВОЙ КОД

    def area(self):
        # ТУТ ТВОЙ КОД

""" # END PROMPT

```

Output:
```
' # BEGIN PROMPT\nclass Rectangle:\n    def __init__(self, width, height):\n        # ТУТ ТВОЙ КОД\n\n    def area(self):\n        # ТУТ ТВОЙ КОД\n\n'
```

```python
grader.check("q2")
```

3.  Напишите класс `BankAccount`, который содержит атрибут `balance` (баланс) и методы `deposit` (для пополнения счета) и `withdraw` (для снятия денег со счета).

```python
# BEGIN SOLUTION NO PROMPT
class BankAccount:
    def __init__(self):
        self.balance = 0

    def deposit(self, amount):
        self.balance += amount
        return f"Баланс увеличен на {amount}. Текущий баланс: {self.balance}."

    def withdraw(self, amount):
        if self.balance >= amount:
            self.balance -= amount
            return f"Сумма {amount} снята со счета. Текущий баланс: {self.balance}."
        else:
            return "Недостаточно средств на счете."


account = BankAccount()
deposit = account.deposit(100)
print(deposit)
withdrawal = account.withdraw(50)
print(withdrawal)


# END SOLUTION
""" # BEGIN PROMPT
class BankAccount:

    # ТУТ ТВОЙ КОД


account = BankAccount()
deposit = account.deposit(100)
print(deposit)
withdrawal = account.withdraw(50)
print(withdrawal)
        # ТУТ ТВОЙ КОД

""" # END PROMPT

```

```python
grader.check("q3")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 2)
```
>>> person = Person("Иван", 25)
>>> 
>>> assert person.age == 25, "Возраст должен быть 25"
>>> assert person.name == "Иван", "Имя должно быть Иван"
>>> assert person.introduce() == "Меня зовут Иван, мне 25 лет.", f"Метод introduce должен возвращать 'Меня зовут Иван, мне 25 лет. Твой метод возвращает: {person.introduce()}'"
```


Tests for q2 (points: 3)
```
>>> 
>>> rectangle = Rectangle(5, 10)
>>> 
>>> assert rectangle.width == 5, "Ширина должна быть 5"
>>> assert rectangle.height == 10, "Высота должна быть 10"
>>> assert rectangle.area() == 50, "Площадь должна быть 50"
```


Tests for q3 (points: 3)
```
>>> 
>>> from utils.hashing import hash_input
>>> 
>>> account = BankAccount()
>>> deposit_operation = account.deposit(100)
>>> withdrawal_operation = account.withdraw(50)
>>> 
>>> assert hash_input(deposit_operation) == "1930cbf854a75446212a6daad8235797e38d1d8665fcfc76de3770bea31b762e", f'{deposit_operation}'
>>> assert hash_input(withdrawal_operation) == "93b5f6565dad47e6a27f799b389d51d33a8794d56657cd9a7456982a44545ae1"
```
